## Continuation Passing Style

Continuation passing style (CPS) makes use of first-class functions by requiring that all functions take an extra argument `k`, which we call the continuation. `k` must be a function, and when a function returns, it calls `k` on the return value and returns that instead.

```scala
// This function takes an integer x and returns x + 1
def addOne(x: Int): Int = {
    val result = x + 1
    result
}
```

```scala
addOne(5)
```

```scala
// In CPS, we would write...
def addOne_cps(x: Int, k: Int => Int): Int = {
    val result = x + 1
    k(result)
}
```

```scala
val identity: Int => Int = y => y
addOne_cps(5, identity)
```

```scala
// Use generics to allow any return type for k
def addTwo_cps[T](x: Int, k: Int => T): T = {
    val result = x + 2
    k(result)
}
```

```scala
addTwo_cps(5, identity)
```

```scala
addTwo_cps(5, x => List(x))
```

## Exercise: Side Effects
Given the following functions, change them to use continuations:

```scala
// Function with side effects
def printFive(f: Int => String): Int = {
    val five = 5
    println(f(five))
    five
}
```

```scala
def printFive_cps(f: Int => String, k: Int => Int): Int = {
    // Your Code
    ???
}
```

```scala
printFive(x => x.toString)
```

```scala
printFive_cps(x => x.toString, y => y * 10)
```

### Exercise: Backtracking
Search a binary tree using CPS. Return true if the tree has a node with the integer `x` as a value.

<!-- grade_id: cell-bdaf875cad5b0678, points: 0, grade: false, solution: true -->
```scala
sealed trait Tree
case object Empty extends Tree
case class Node(left: Tree, value: Int, right: Tree) extends Tree

def search(t: Tree, x: Int, fail_continuation: () => Boolean): Boolean =
    // Your Code
    ???

```

<!-- grade_id: cell-df2188387bae0d26, points: 1, grade: true, solution: false -->
```scala
val t = Node(Empty, 10, Node(Empty, 6, Empty))
assert(search(t, 10, () => false))
assert(!search(t, 0, () => false))
```

```scala

```